# Anhalt
![](image/logo.png)
<div class="alert alert-block alert-info">
</div>




## Intoduction to Probability

To do as per slides

## <center> Grundlagen der Wahrscheinlichkeitsrechnung </center><br>
<center> Was ist Wahrscheinlichkeit ? </center><br>
<center> P.G.Plöger </center><br>

<center>Quellen / Zitate</center>
<table><tr><td><img src="image/image1.png" width="200"></td><td><img src="image/image2.png" width="200"></td></tr></table>

* Wie wahrscheinlich ist eine „6“ beim nächsten Wurf? <br>
* Wenn Sie eine E-Mail erhalten, die die Wörter "frei", "Geld", "billig", "Kredit" enthält, mit welcher Sicherheit wird es in den Spam-Ordner gefiltert? => Machine Learning <br>
* Wird das Pferd No. 8 gewinnen? <br>
* Morgen regnet es mit 80% Wahrscheinlichk. <br>
* Steigen die Börsenkurse morgen? <br>
* Wie wahrscheinlich sind Aliens im Universum? <br>

## <center>Probability Theory </center>

<center>
<img src="image/image3.png"   width="500" >
</center>


All diese Fragen werden unter einen Hut gebracht (erst 1931 (de Finetti), dann 1933 (Kolmogoroff), dann 
1946 (Cox)) <br>
.. und das hat alles bloß ca 300 Jahre gedauert.. <br>

## <center> Outline </center></h1>
* <b> Studiere Modell Fälle <br>
* Formalisiere (etwas....) <br> </b>
* Bilde Begriffe <br>
* Bilde Axiome (Kolmogoroff im Original) <br>
* Zeige: die funktionieren (einfache Folgen..) <br>
* Probleme und Erweiterungen <br>

## <center> Mengen + Venn Diagr. </center>
<center>
<img src="image/image4.png"  width="800">
</center>

## <center> Was haben diese „Modelle“ gemeinsam? </center></h1>

<center>
<img src="image/image16.png"  width="500">
</center>

## <center> Übersetzungen: {} -- Events </center>
<center>
<img src="image/image5.png" width="600" >
</center>


<center>
<img src="image/image6.png" width="600">
</center>

## <center> W-Generator Schema </center>
<center>
<img src="image/image17.png" width="600">
</center>

## <center> Outline </center>
* Studiere Modell Fälle
* Formalisiere (etwas....)
* <b> Bilde Begriffe </b>
* Bilde Axiome (Kolmogoroff im Original)
* Zeige: die funktionieren (einfache Folgen..)
* Probleme und Erweiterungen

## <center> Begriffe der Wahrscheinlickeit </center>
Ein <b>Ereignis</b> ist der Ausgang / das Resultat eines Experiments. <br>
Es kann als Menge A, B, O, I, E_i .... dargestellt werden. <br>

Ein <b>Elementar Ereignis</b> (oder Atom) E_i ist ein  
unteilbares Ereignis, d.h. es gibt KEINE Mengen A, B mit Ei=A+B und Ei≠A und Ei≠B).  <br>


## <center> Begriffe der Wahrscheinlickeit  </center>

Der <b>Ereignisraum</b>  enthält die Menge aller möglichen Elementar Ereignisse.<br>
Die Potenzmenge von Ω enthält 
dann alle Vereinigungsmengen,Komplemente und Schnittmengen der Elementar-ereignisse und bildet damit die sogenannte <b>Ereignis-Algebra</b> 𝕱. <br>
Sie enhält ALLE Ereignisse. <br>
Zwei Spezial-Ereignisse: <br>
Ω nennt man auch das <b>sichere Ereignis</b>, ∅ das <b>unmögliche Ereignis.</b>


## <center> Darstellungen  </center>

<center>
<img src="image/image8.png" width="500" >
</center>

* Geordnete Paare (i,j) für zwei 4-Seiten-Würfel
* der Baum auf der rechten Seite die spiegelt den sequenzieller Charakter des Experiments.
* graues Ereignis: {(1, 1), (1, 2), (1, 3), (1, 4)}
* 1 Knoten des Baums == 1 Ereignis == Menge aller Blätter, die diesem Knoten nachgeschaltet sind

## <center> Bäume für allgemeine A, B  </center>
<center>
<img src="image/image18.png" width="300" >
</center>
<br>

Gestufter Verlauf: <br>
erst A dann B <br>
An den Blättern: <br>
* alle Kombinationen <br>
* Das Ereignis steckt <br>
im Pfad zum Blatt <br>
(A,B) oder (A,$B^c$) etc. 

## <center> W-Generator Schema  </center>
<center>
<img src="image/image19.png" width="500" >
</center>
<br>
Beispiele für Ω: beim Würfeln: <br>
Seitenflächen eines Würfels => {1,2,3,4,5,6} <br>
Seitenflächen zweier Würfel => {(1,1),(1,2),(2,1),...,(6,6)} <br>
Summe zweier Würfel => {2,3,...,7,..., 12} <br>

## <center> Elemente eines probabilistischen Models </center>

1. die Wahrscheinlichkeitsfunktion, die jedem Ereignis A ∈  𝕱 eine nichtnegative Zahl P(A) zuordnet, 
=> genannt die/das Wahrscheinlichkeit(smaß) von A. 

2. P(A) bündelt unser Wissen oder unsere Überzeugungen über die "Wahrscheinlichkeit" des Ereignisses A. 

Eine Wahrscheinlichkeitsfunktion muss bestimmten Axiomen genügen.

## <center> Outline </center>

* Studiere Modell Fälle
* Formalisiere (etwas....)
* Bilde Begriffe
* **Bilde Axiome (Kolmogoroff im Original)**
* Zeige: die funktionieren (einfache Folgen..)
* Probleme und Erweiterungen

## <center> Axiome der Wahrscheinlickeit (Kolmogoroff 1933) (Zitat) </center>


III. $0 \leqslant P(E_i)$

IV. $P(E_i) =1 $

V. If A$\cap$B = ∅ then P(A$\cup$B) = P(A)+ P(B)


NB: die Kolmogoroff Original-Axiome I. und II. beziehen sich auf die Ereignis-Algebra <br>
NB: Im Fall V. nennt man A und B disjunkt

## <center> Die Ereignis Algebra 1 </center>
<b>Begriffe: </b> <br>

<b>Experiment: </b>  werfen einer Münze <br>
<b>Ereignis: </b>  H, T (elementary events) <br>
<b>Algebra: </b>  {H,T},{H},{T},∅ <br>
<b>Wahrscheinlk.: </b>  P({H})= P({T})=1/2, <br>
P({H,T})=1, <br>
P(∅)see below… <br>
Bemerkung: die korrekte Notation P({H}) wird
fast immer durch die „schlampigere“ Schreibweise
P(H) ersetzt!

## <center> Die Ereignis Algebra 2 </center>
<b>Begriffe: </b> <br>
<b>Experiment: </b> toss two 4-sided dice. <br>
<b>Elm. Ereigniss: </b> (1,1),(1,2),…(4,4) <br>
<b>Ereignisse: </b> {(i,j)| i=j },<br>
{(i,j)| i=4 ∨ j=4},
<center>
<img src="image/image20.png" width="300" >
</center>

## <center> Outline </center>
* Studiere Modell Fälle
* Formalisiere (etwas....)
* Bilde Begriffe
* Bilde Axiome (Kolmogoroff im Original)
* **Zeige: die funktionieren (einfache Folgen..)**
* Probleme und Erweiterungen

## <center> Einige direkte Folgerungen </center>

<center>
<img src="image/image21.png" width="500" >
</center>

## <center> Nutze die Axiome für weitere einfache Folgerugen </center>
<center>
<img src="image/image22.png" width="500" >
</center>
<br>
NB: diese „harmlose“ Regel ist ein häufig verwendeter <br>
Rechentrick: berechne statt der direkten Wahrscheinlichkeit <br>
die Wahrscheinlichkeit des Komplements <br>

## <center> Example </center>
<center>
<img src="image/image23.png" width="500" >
</center>
<br>
* Man kann das Wahrscheinlichkeitsmaß gut durch die Fläche von Mengen veranschaulichen.
* Wie hängen Mengenoperationen mit diesen Flächenmassen zusammen?
* Beachte: im obigen Beispiel bezeichnet A zweierlei:

1. die Menge A und 
2. das Ereignis, das ein Element aus A auswählt wurde

## <center> Outline </center>
* Studiere Modell Fälle
* Formalisiere (etwas....)
* Bilde Begriffe
* Bilde Axiome (Kolmogoroff im Original)
* Zeige: die funktionieren (einfache
* Folgen..)
* Probleme und Erweiterungen

## <center>Sichten auf Ungewissheit und Wahrscheinlichkeit  </center>

Ziel der Wahrscheinlichkeitstheorie ist es, die Unsicherheit zu beschreiben.
Quellen / Interpretationen für die Unsicherheit sind:

1. <b> Frequentivistisch: </b> <br>
Wahrscheinlichkeiten sind relative Häufigkeiten.
(z. B. die relative Häufigkeit des Kopfwurfs)

2. <b> Objektivistisch: </b>  <br>
Wahrscheinlichkeiten sind reale Aspekte der Welt.
(z. B. die Wahrscheinlichkeit, dass ein Atom in der nächsten Stunde zerfällt)

3. <b> Subjektivistisch: </b> <br>
Wahrscheinlichkeiten beschreiben den Grad des Glaubens eines Agenten,
(z. B. ist es plausibel, dass Außerirdische existieren)

## <center>Frequentivistische Sicht: Häufigkeiten   </center>

Der Frequentivist interpretiert Wahrscheinlichkeiten als relative Häufigkeiten. <br>
 
* Wenn in einer Folge von n unabhängigen, identisch verteilten (i.i.d.) Experimenten (Versuchen) ein Ereignis k(n) mal auftritt, 
dann ist die relative Häufigkeit des Ereignisses gegeben durch k(n)/n.
also

<center>
<img src="image/image24.png" width="500" >
</center>
<br>

* die Wahrscheinlichkeit wird dann definiert als 
der Grenzwert der relativen Häufigkeiten:
<center>
<img src="image/image25.png" width="50" >
</center>
<br>
* z.B. ist die Wahrscheinlichkeit des Ereignisses „H“ in einer Reihenfolge des wiederholten Werfens einer fairen Münze gleich ½.

* Die Frequentivistische Position ist am einfachsten zu verstehen, sie hat aber mehrere signifikante Mängel:
     * die Definition ist zirkulär => see image
     * sie ist beschränkt auf den Fall i.i.d,
     * das Problem der Referenz Klassen

<h1><center>Was heißt hier Konvergenz (?) </center></h1>
<center>
 <img src="image/image10.png" width="500" >
</center>



```python
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
data = np.random.rand(1000)
ax = plt.hist(data)
data = np.random.rand(1000)
A = np.random.rand(12,1000)
#print(A)
B = np.mean(A,0)
#print(B)
plt.hist(B,bins=20000)
plt.show()
```

<h1><center>The Reference Class problem </center></h1>
Sortierkriterium zu schwach  <br>
=> zu viele Proben in jeder Klasse  <br>
=> zu wenig Selektivität <br>

Sortierkriterium zu stark  <br>
=> zu wenig Proben in jeder Klasse, eine viel zu hohe Anzahl von Klassen  <br>
=> zu hohe Selektivität <br>

Was ist die "richtige" Klassengröße, <br>
um eine frequentistische Sicht der Wahrscheinlichkeit zu unterstützen? <br>

Eine offene Frage ….<br>

<h1><center>Objektivistische Interpretation: nicht deterministische Eigenschaft </center></h1>

Für den Objektivisten sind Wahrscheinlichkeiten reale Aspekte der Welt

Das Ergebnis einer Beobachtung oder eines Experiments ist nicht deterministisch, sondern beinhaltet physikalische Zufallsprozesse

Beispiel: radioaktiver Zerfall

<h1><center>Subjektivistische Interpretation: Glaubhaftigkeit </center></h1>

Der Subjektivist verwendet Wahrscheinlichkeiten, um den Grad des Glaubens eines Agenten an etwas zu charakterisieren, anstatt einen physikalischen zufälligen Prozess zu charakterisieren <br>

Dies ist die relevanteste Interpretation von Wahrscheinlichkeiten in der Künstlichen Intelligenz. <br>

Wir definieren die Plausibilität eines Ereignisses als den Grad des Glaubens an das Ereignis oder die subjektive Wahrscheinlichkeit des Ereignisses <br>

Es ist natürlich anzunehmen, dass Plausibilität / Glauben Bel (* | *) durch reelle Zahlen repräsentiert werden kann, dass die Regeln qualitativ dem gesunden Menschenverstand entsprechen und dass die Regeln mathematisch konsistent sind <br>


<h1><center>Subjektivistische Interpretation: Wie „misst“ man Glaubhaftigkeit? </center></h1>

Habe Sie eine Idee? <br>

Denken Sie an die Börse... <br>

Die Plausibilität eines Ereignisses wird angenährt durch die Menge der Leute, die das Eintreten des Ereignis plausibel finden und die darauf wetten würden (plausibles Schließen). <br>

Dieser „Mechanismus“ ersetzt das Experiment bei Kolmogoroff.  <br>

Mit dieser Methode können wir den ORANGEN Fragen auf Folie 2 eine Wahrscheinlichkeit zuordnen obwohl es keine Wiederholbarkeit gib. <br>

PROBLEM: <br>
Gibt es jetzt MEHRERE Wahrscheinlichkeit ?!?

<h1><center>Degrees of Belief </center></h1>

(Richard T.) Cox's theorem (1946): <br>
     
Bel(x|A) is isomorphic to  a probability measure p(x|x) that  satisfies the axioms of (objective) probabilities.

Die „Belief“ Funktion gehorcht den selben Gesetzen wie die Wahrscheinlichsfunktion 


<h1><center>Resümee der Sichten auf Wahrscheinlichkeit </center></h1>

* Wahrscheinlichkeit als Gewinnwahrscheinlichkeit
     * eine relative Häufigkeit: beobachtete Gewinne im Vergleich zu gesamten Studien.<br>  
     <b>Frequentivistische Sicht </b>
    
* Wahrscheinlichkeit als Attribut von Objekten:
    * eine faire Münze: Chancen 50/50 sind "eingebaut", d. h. 	physische Eigenschaft der Objekte <br> <b>Objektivisten-Ansicht </b> 

* Wahrscheinlichkeit als Unsicherheit:
	* der Plausible Zustand des Agenten: 80% Chance auf 	Hohlraum, wenn es bereits Zahnschmerzen gibt  <br> <b>subjektivistische Sichtweise</b>

<h1><center>Zusammenfassung </center></h1>
Recap über Mengen
* Kolomoroff‘s Axiome: 
    * Ereignisse sind wie Mengen
* Ereignis Algebra: Teilmengen
* Kartesisches Produkt / Bäume:
    * Mehrstufige Experimente
* Unterschiedliche Interpretationen der Wahrscheinlichkeit:
    * Frequentivistisch, Objektiv, Subjektiv

# Klassenübung

<div class="alert alert-block alert-info">
</div>




## Introduction to Python

Copyright [Steve Phelps](http://sphelps.net) 2014


## Python is interpreted

- Python is an _interpreted_ language, in contrast to Java and C which are compiled languages.

- This means we can type statements into the interpreter and they are executed immediately.


```python
5 + 5
```

```python
x = 5
y = 'Hello There'
z = 10.5
```

```python
x + 5
```


# Assignments versus equations

- In Python when we write `x = 5` this means something different from an equation $x=5$.

- Unlike variables in mathematical models, variables in Python can refer to different things as more statements are interpreted.

```python
x = 1
print('The value of x is ', x)

x = 2.5
print('Now the value of x is ', x)

x = 'hello there'
print('Now it is ', x)
```

## Calling Functions

We can call functions in a conventional way using round brackets

```python
import math
print(round(math.pi,2))
```

## Types

- Values in Python have an associated _type_.

- If we combine types incorrectly we get an error.

```python
print(y)
```

```python
y + 5
```

## The type function

- We can query the type of a value using the `type` function.

```python
type(1)
```

```python
type('hello')
```

```python
type(2.5)
```

```python
type(True)
```

## Null values

- Sometimes we represent "no data" or "not applicable".  

- In Python we use the special value `None`.

- This corresponds to `Null` in Java or SQL.

```python
result = None
```

- When we fetch the value `None` in the interactive interpreter, no result is printed out.


```python
result
```

- We can check whether there is a result or not using the `is` operator:

```python
result is None
```

```python
x = 5
x is None
```

## Converting values between types

- We can convert values between different types.

- To convert an integer to a floating-point number use the `float()` function.
- To convert a floating-point to an integer use the `int()` function.

```python
x = 1
print(type(x))
print(x)

```

```python
y = float(x)
print(type(y))
print(y)
```

```python
print(int(y))
```

## Converting to and from ASCII values

- The functions `chr()` and `ord()` can be used to convert characters from and to [ASCII](https://en.wikipedia.org/wiki/ASCII).

```python
print(ord('a'))
```

```python
print(chr(97))
```

## Variables are not typed

```python
y = 'hello'
print('The type of the value referred to by y is ', type(y))
y = 5.0
print('And now the type of the value is ', type(y))
```

## Polymorphism
- The meaning of an operator depends on the types we are applying it to.


```python
1//5
```

```python
1.0 / 5.0
```

```python
1 + 1
```

```python
'a' + 'b'
```

```python
'1' + '1'
```

## Conditional Statements and Indentation

- The syntax for control structures in Python use _colons_ and _indentation_.

- Beware that white-space affects the semantics of Python code.

```python
x = 5
if x > 0:
    print('x is strictly positive')
    print(x)
    
print('finished.')
```

```python
x = 0
if x > 0:
    print('x is strictly positive')
print(x)
    
print('finished')
```

## Lists


We can use _lists_ to hold an ordered sequence of values.

```python
l = ['first', 'second', 'third']
print(l)
```

```python
another_list = ['first', 'second', 'third', 1, 2, 3]
print(another_list)
```

## Mutable Datastructures

Lists are _mutable_; their contents can change as more statements are interpreted.

```python
l.append('fourth')
print(l)
```

## References

- Whenever we bind a variable to a value in Python we create a *reference*.

- A reference is distinct from the value that it refers to.

- Variables are names for references.

```python
X = [1, 2, 3]
Y = X
```

- The above code creates two different references (named `X` and `Y`) to the *same* value `[1, 2, 3]`

- Because lists are mutable, changing them can have side-effects on other variables.

- If we append something to `X` what will happen to `Y`?

```python
X.append(4)
X
```

```python
Y
```

## State and identity

- The state referred to by a variable is *different* from its identity.

- To compare *state* use the `==` operator.

- To compare *identity* use the `is` operator.

- When we compare identity we check equality of references.

- When we compare state we check equality of values.


```python
X = [1, 2]
Y = [1]
Y.append(2)
```

```python
X == Y
```

```python
X is Y
```

```python
Y.append(3)
X

```

```python
X = Y
```

```python
X is Y
```

## Iteration

- We can iterate over each element of a list in turn using a `for` loop:


```python
for i in l:
    print(i)
```

- To perform a statement a certain number of times, we can iterate over a list of the required size.

```python
for i in [0, 1, 2, 3]:
    print("Hello!")
```

## For loops with the range function

- To save from having to manually write the numbers out, we can use the function `range()` to count for us.  As in Java and C, we count starting at 0.

```python
range(4)
```

```python
for i in range(4):
    print("Hello!")
```

## List Indexing

- Lists can be indexed using square brackets to retrieve the element stored in a particular position.



```python
print(l[0])
```

```python
print(l[0])
```

## List Slicing

- We can also a specify a _range_ of positions.  

- This is called _slicing_.

- The example below indexes from position 0 (inclusive) to 2 (exclusive).


```python
print(l[0:2])
```

```python
print(l[:2])
```

- If we leave out the final index it implies the end of the list:

```python
print (l[2:])
```

## Negative Indexing

- Negative indices count from the end of the list:


```python
print(l[-1])
```

```python
print(l[:-1])
```

## Collections

- Lists are an example of a *collection*.

- A collection is a type of value that can contain other values.

- There are other collection types in Python:

    - `tuple`
    - `set`
    - `dict`

## Tuples

- Tuples are another way to combine different values.

- The combined values can be of different types.

- Like lists, they have a well-defined ordering and can be indexed.

- To create a tuple in Python, use round brackets instead of square brackets

```python
tuple1 = (50, 'hello')
tuple1
```

```python
tuple1[0]
```

```python
type(tuple1)
```

## Tuples are immutable

- Unlike lists, tuples are *immutable*.  Once we have created a tuple we cannot add values to it.


```python
tuple1.append(2)
```

## Sets

- Lists can contain duplicate values.

- A set, in contrast, contains no duplicates.

- Sets can be created from lists using the `set()` function.

```python
X = set([1, 2, 3, 3, 4])
X
```

```python
type(X)
```

- Alternatively we can write a set literal using the `{` and `}` brackets.

```python

X = {1, 2, 3, 4}
type(X)
```

## Sets are mutable

- Sets are mutable like lists:

```python
X.add(5)
X
```

```python
X.add(5)
X

```

## Sets are unordered

- Sets do not have an ordering.

- Therefore we cannot index or slice them:

```python
X[0]
```

## Operations on sets

- Union: $X \cup Y$


```python
X = {1, 2, 3}
Y = {4, 5, 6}
X.union(Y)
```

```python
X = {1, 2, 3, 4}
Y = {3, 4, 5}
X.intersection(Y)
```

- Difference $X - Y$:


```python
X - Y
```

## Arrays

- Python also has fixed-length arrays which contain a single type of value

- i.e. we cannot have different types of value within the same array.   

- Arrays are provided by a separate _module_ called numpy.  Modules correspond to packages in e.g. Java.

- We can import the module and then give it a shorter _alias_.


```python
import numpy as np
```

- We can now use the functions defined in this package by prefixing them with `np`.  

- The function `array()` creates an array given a list.

```python
x = np.array([0, 1, 2, 3, 4])
print(x)
print(type(x))
```

## Functions over arrays

- When we use arithmetic operators on arrays, we create a new array with the result of applying the operator to each element.

```python
y = x * 2
print(y)
```

- The same goes for functions:

```python
x = np.array([-1, 2, 3, -4])
y = abs(x)
print(y)
```

## Populating Arrays

- To populate an array with a range of values we use the `np.arange()` function:


```python
x = np.arange(0, 10)
print(x)
```

- We can also use floating point increments.

```python
x = np.arange(0, 1, 0.1)
print(x)
```

## Basic Plotting

- We will use a module called `matplotlib` to plot some simple graphs.

- This module provides func

```python
import matplotlib.pyplot as plt
%matplotlib inline

y = x*2 + 5
plt.plot(x, y)

```

## Plotting a sine curve

```python
from numpy import pi, sin

x = np.arange(0, 2*pi, 0.01)
y = sin(x)
plt.plot(x, y)
```

## Plotting a histogram

- We can use the `hist()` function in `matplotlib` to plot a histogram

```python
# Generate some random data
data = np.random.randn(1000)

ax = plt.hist(data)
```

## Computing histograms as matrices

- The function `histogram()` in the `numpy` module will count frequencies into bins and return the result as a 2-dimensional array.

```python
np.histogram(data)
```

## Defining new functions



```python
def squared(x):
    return x ** 2

print(squared(5))
```

## Local Variables

- Variables created inside functions are _local_ to that function.

- They are not accessable to code outside of that function.

```python
def squared(x):
    result = x ** 2
    return result

print(squared(5))
```

```python
print(result)
```

## Functional Programming

- Functions are first-class citizens in Python.

- They can be passed around just like any other value.



```python
print(squared)
```

```python
y = squared
print(y)
```

```python
print(y(5))
```

## Mapping the elements of a collection

- We can apply a function to each element of a collection using the built-in function `map()`.

- This will work with any collection: list, set, tuple or string.

- This will take as an argument _another function_, and the list we want to apply it to.

- It will return the results of applying the function, as a list.

```python
map(squared, [1, 2, 3, 4])
```

- this was changed in phyton 3 ```map``` only returns iterator to some pointers to objects

- thus we need to put a "*" and construct a list from it, which is then printed

```python
[*map(squared, [1, 2, 3, 4])]
```

## List Comprehensions

- Because this is such a common operation, Python has a special syntax to do the same thing, called a _list comprehension_.


```python
[squared(i) for i in [1, 2, 3, 4]]
```

## Cartesian product using list comprehensions



<img src="files/220px-Cartesian_Product_qtl1.svg.png">

The [Cartesian product](https://en.wikipedia.org/wiki/Cartesian_product) of two collections $X = A \times B$ can be expressed by using multiple `for` statements in a comprehension.


```python
A = {'x', 'y', 'z'}
B = {1, 2, 3}
{(a,b) for a in A for b in B}
```

## Filtering using a list comprehension

- Again, because this is such a common operation, we can use simpler syntax to say the same thing.

- We can express a filter using a list-comprehension by using the keyword `if`:

```python
data = [-5, 2, 3, -10, 0, 1]
[x for x in data if x > 0]
```

- We can also filter and then map in the same expression:

```python
from numpy import sqrt
[sqrt(x) for x in data if x > 0]
```

## The reduce function

- The `reduce()` function recursively applies another function to pairs of values over the entire list, resulting in a _single_ return value.

```python
from functools import reduce
```

```python
reduce(lambda x, y: x + y, [0, 1, 2, 3, 4, 5])
```

## Reading Text Files

- To read an entire text file as a list of lines use the `readlines()` method of a file object.


```python
f = open('/etc/group')
result = f.readlines()
f.close()

```

```python
# Print the first line
print(result[4])
```

To concatenate into a single string:


```python
single_string = ''.join(result)
```

```python
print(single_string)
```

# Prof Ploeger and TA please fix this

# Aufgabe

<div class="alert alert-block alert-info">
</div>




<div class="alert alert-block alert-danger">

### Student Nutzername 1: ---------

### Student Nutzername 2: ---------

</div>


<div class="alert alert-block alert-warning">
<b>Example:</b> Indem Sie hier Ihren Namen schreiben, stimmen Sie zu, dass die folgende Aufgabe von mir / uns erledigt wurde.
</div>

## Question 1 : (10	pts)
<b> Assume </b>	:	you are	the	TEACHER	of	this lecture and you want to check,	how	good the students understood the last lecture. Formulate 5 questions and 5 expected	answers	by which you check how well the	students did understand	the	 lecture.		<br>
In each cell below type one question and answer

### Question and Answer 1

<!-- grade_id: cell-9149c479d159bb6b, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

### Question and Answer 2

<!-- grade_id: cell-45b81cf514710b89, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

### Question and Answer 3

<!-- grade_id: cell-28f499ed4e0d0e7a, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

### Question and Answer 4

<!-- grade_id: cell-dc0c5eb1f986c75b, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

### Question and Answer 5

<!-- grade_id: cell-9bf3521b9a5435ef, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2 : (10 pts)
Out of the students in a class, 60% are geniuses, 70% love chocolate, and 40% fall into both categories. Determine the probability that a randomly selected student is neither a genius nor a chocolate lover 

<!-- grade_id: cell-0665a49a5aa9268c, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

let <br>
xRB be the tuple with the name and probability of RottenBeans <br> 
xF be the tuple with the name and probability of fruits <br>
Which have been declared below

```python
xRB = ('RottenBean',0.001); #MY NEW data type prob.boolean vars: tuple(value, probability)
xF = ('FruitsOK',0.9)

print(xRB)
print(xF)
```

CHANGE the implementation of the Boolean variables: Use the class Fraction instead of float numbers to represent probabilities (hint: „from fractions import Fraction“) <br>
Please use the same variable name for creation fraction probabilities(xRB and xF)

Solution: 

Follow the hint: include some DT which can do fractions. You do this via importing the module in the current name space.

```python
from fractions import Fraction
```



Check what this class can do for you:

HINT:: scroll almost to the end FIRST to check the EXAMPLES!

```python
help(Fraction)
```

```python
Fraction(1,1000)
```

```python
myTemp = 1000*Fraction(1,1000)
```

```python
sqrt(myTemp)
```

Oops we need the math module,

```python
from math import *
```

When we do the import in the style above (i.e. using "*" at the end), we can use ALL functions without a prefix like "math.sqrt" etc. This is usally NOT recommended in order to avoid cluttering the name space of the runtime system.

```python
sqrt(myTemp)
```

```python
myTemp = sqrt(myTemp)
```

```python
type(myTemp)
```

```python
my49 = Fraction(Fraction(1,7),7);print(my49)
```

```python
Fraction(Fraction(1,7),Fraction(7,1))
```

```python
myTemp = Fraction(Fraction(1,7),Fraction(1,7))
```

```python
type(myTemp)
```

```python
sqrt(my49)
```

```python
Fraction(1,0)
```

```python
Fraction(0,1)
```

```python
print(pi)
print(type(pi))
Fraction(pi)
```

```python
print('{0:12.40g}'.format(pi-Fraction(pi)))
```

```python
884279719003555 / 281474976710656
```

## Question 3 : (10 pts)

CHANGE the implementation of the Boolean variables: Use the class Fraction instead of float numbers to represent probabilities (hint: „from fractions import Fraction“)
Please use the same variable name for creation fraction probabilities(xRB and xF)

<!-- grade_id: cell-4a4420e4c80f065f, points: 0, grade: false, solution: true -->
```python
from fractions import Fraction
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f4aca995aca36634, points: 10, grade: true, solution: false -->
```python
# This cell is for grading purpose . Please ignore it. 

```

The following code returns the complement of the input boolean variable

```python
def myCOMPLEMENT(myVar):
    '''Returns the compement of the boolean variable

    Keyword arguments:
    mVar -- tuple : (Name, prob.)
    
    returns (Name, 1-prob)
    '''
    
    return (myVar[0],1-myVar[1])
```

```python
myNew = myCOMPLEMENT(xRB)
myNew
```

Change the implementation of the function "myCOMPLEMENT “: <br>
when calling, ADD  the Letter "N_" in front of the new variable name AND if there are two "NN_" at the beginning DROP this prefix 

Strings in python are delimited by " or ' or ```, like in many other languages the behave like arrays.

So we can access each letter by indexing.

```python
myTest = 'ABC'
```

```python
myTest[0]
```

What methods are offered by data tyoe string? (follow the hint: search for startswith)

```python
help(type(myTest))
```

```python
print("Text start with 'AB'? : ", myTest.startswith('AB'))
print("Text start with 'ABC'? : ", myTest.startswith('ABC'))
print("Text start with 'C'? : ", myTest.startswith('C'))
print("Text start with ''? : ", myTest.startswith(''))
```

## Question 4 : (20 pts)

Change the implementation of the function "myCOMPLEMENT “: 
when calling, ADD the Letter "N_“ in front of the new variable name AND if there are two "NN_“ at the beginning DROP this prefix

<!-- grade_id: cell-2f42843bcb2e9eda, points: 0, grade: false, solution: true -->
```python
def myCOMPLEMENT(myVar):
    '''Returns the compement of the boolean variable

    Keyword arguments:
    mVar -- tuple : (Name, prob.)
    
    returns (Name, 1-prob)
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-022ee12c1e35f2c9, points: 20, grade: true, solution: false -->
```python
assert myCOMPLEMENT(('xx',Fraction(1,10))) == ('N_xx', Fraction(9,10))
assert myCOMPLEMENT(('NN_xx',Fraction(9,10))) == ('xx', Fraction(1,10))
```

Run the following code for AND and OR operators. <br>


```python
def myAND(myVar1,myVar2):
    '''Returns a probalistic AND on two boolean variable

    Keyword arguments:
    myVar1,myVar2 -- two tuples : (Name, prob.)
    
    returns:: the probabilty tuple of both being true (on average)
    '''
    
    return (myVar1[0] + '_A_'+ myVar2[0],myVar1[1]*myVar2[1])

def myOR(myVar1,myVar2):
    '''Returns a probalistic OR of two boolean variable

    Keyword arguments:
    myVar1,myVar2 -- two tuples : (Name, prob.)
    
    returns:: the probabilty tuple of ONE of them being true
    '''
    return (myVar1[0] + '_O_' + myVar2[0],myVar1[1] + myVar2[1] - myVar1[1]*myVar2[1])

```

```python
mytemp = myAND(xRB,xF)
mytemp
```

```python
mytemp = myOR(xRB,xF)
mytemp
```

## Question 5 :  (20 pts)
Use the functions myAND and myOR to Add 3 more Boolean operators (e.g. NOR, NAND etc...)

<!-- grade_id: cell-0a0362507bbbbf42, points: 0, grade: false, solution: true -->
```python
def myNAND(myVar1,myVar2):
    '''Returns a probalistic NAND on two boolean variable

    Keyword arguments:
    myVar1,myVar2 -- two tuples : (Name, prob.)
    
    returns:: the probabilty tuple of both being true (on average)
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1a955bfdabd96ae8, points: 0, grade: false, solution: true -->
```python
def myNOR(myVar1,myVar2):
    '''Returns a probalistic NOR of two boolean variable

    Keyword arguments:
    myVar1,myVar2 -- two tuples : (Name, prob.)
    
    returns:: the probabilty tuple of ONE of them being true
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-fc250c8fb83cce25, points: 0, grade: false, solution: true -->
```python
def myXOR(myVar1,myVar2):
    '''Returns a probalistic XOR of two boolean variable

    Keyword arguments:
    myVar1,myVar2 -- two tuples : (Name, prob.)
    
    returns:: the probabilty tuple of ONE of them being true
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f0c95d4d5b074ecf, points: 7, grade: true, solution: false -->
```python
# This cell is for grading purpose . Please ignore it.

```

<!-- grade_id: cell-76a02f2a14147af7, points: 7, grade: true, solution: false -->
```python
# This cell is for grading purpose . Please ignore it.

```

<!-- grade_id: cell-3fd87f0249fd7aec, points: 6, grade: true, solution: false -->
```python
# This cell is for grading purpose . Please ignore it.

```

## Question 6 : (20 pts)

Assume the lines in diagram below are pipes and the boxes are radiators, which block with a certain likelihood.
We want that water flows from a to b. Assuming that two radiators are connected in sequence then the water will flow with P(water_flow) = myAND(radiator_1_OK,radiator_2_OK). When twoare connected in parallel
P(water_flow) = myOR(radiator_1_OK,radiator_2_OK).

Hint: Apply the functions ftom exercise 4. 

Assume: P_K1 = 0.99, P_K2 = 0.95, P_K3 = 0.05, P_K4 = 0.56

How probable is there a connection
from (a) to (b) ??  
<img src="image/image_circuit.png" />

<!-- grade_id: cell-e0f544b75362aa49, points: 0, grade: false, solution: true -->
```python
def is_connection(myK1, myK2, myK3, myK4):
    # YOUR CODE HERE
    raise NotImplementedError()
    return myAll
```

```python
myK1 = ('K1',1-Fraction(1,100))
myK2 = ('K2',1-Fraction(5,100))
myK3 = ('K3',Fraction(5,100))
myK4 = ('K4',Fraction(56,100))
is_connection(myK1, myK2, myK3, myK4)
```

<!-- grade_id: cell-01ebc3a722493801, points: 20, grade: true, solution: false -->
```python
# This cell is for grading purpose . Please ignore it.

```

## Question 7: (120 pts)
Program a die as a function call :: <br>
<b> myDice(NumFaces,repeatCount) </b> <br>
where NumFaces is 4,6,8,10,12,20, returning a tuple (or array) of length repeatCount containig the sequence of outcomes. <br>
a) print a statistic: for <b>NumFaces</b> = 6 and <b>repeatCount</b> =1200 all numbers should appear roughly 1/6 of the times <br>
b) extend the experiment: use N=2,3,4,5,6 many dice and define as the outcome the SUM of all faces (e.g. for N=2 the sample space is now {2,3,..., 11,12} <br>
c) build some new statistics: how likely are all possible outcomes when using N dice? <br>
d) draw the statistics in a diagram, compare the functions for different N <br> 
e) can you find a relation between the outcomes? Hint: Wiki Faltung

```python
import numpy as np
from numpy.random import randint
```

```python
help(randint)
```

```python
randint(2, size=10)
```

```python
a = randint(1,6, size=20)
help(np.mean)
```

```python
list(a).count(1)
```

<!-- grade_id: cell-ef981b32ecc9666b, points: 30, grade: true, solution: true -->
```python
def myDice(numFaces, repeatCount):
    # YOUR CODE HERE
    raise NotImplementedError()
    
myThrows = myDice(6,1200)
myThrows
```

<!-- grade_id: cell-17249a37c7114a4b, points: 30, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-39f4fee0878ffc4c, points: 20, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e69c942ba26b422a, points: 20, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a6c5e971c099fd9c, points: 20, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Question 8 (20 pts)
<b> Prove the following Kolmogorov axiom </b> <br>
P(A$\cup$B) = P(A) + P(B) – P(A $\cap$ B) <br> 
HINT: in ipython markdown use the terms   " \cup "  for union($\cup$) symbol and and "\cap“  for intersection($\cap$) symbol.  


<!-- grade_id: cell-d3f5a3e62996a1b8, points: 20, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 9: (20 pts)
A six-sided die is loaded in a way that each even face is twice as likely as each odd face. All even faces are equally likely, as are all odd faces. Construct a probabilistic model for a single roll of this die and find the probability that the outcome is less than 4

<!-- grade_id: cell-0d2c4ff1aa294208, points: 20, grade: true, solution: true -->
YOUR ANSWER HERE